Atypical Biological Motion Perception in Children with Attention Deficit Hyperactivity Disorder: Dissociating the Roles of Local Motion and Global Configuration Processing

Author Response

The following is the authors’ response to the original reviews.

eLife assessment

This manuscript represents a cleanly designed experiment for assessing biological motion processing in children (mean age = 9) with and without ADHD. The group differences concerning accuracy in global and local motion processing abilities are solid, but the analyses suggesting dissociable relationships between global and local processing and social skills, age, and IQ need further interrogation. The results are useful in terms of understanding ADHD and the ontogenesis of different components of the processing of biological motion.

We thank the editors for the positive assessment of our manuscript. We have carefully considered the reviewers’ constructive and helpful comments and revised our manuscript accordingly. To address the question about the dissociable relationships between global and local BM processing, we have provided more evidence and additional analyses in this revised version.

Reviewer #1 (Public Review):

Summary:

The paper presents a nice study investigating differences in biological motion perception in participants with ADHD in comparison with controls. Motivated by the idea that there is a relationship between biological motion perception and social capabilities, the authors investigated local and global (holistic) biological motion perception, the group, and several additional behavioral variables that are affected in ADHS (IQ, social responsiveness, and attention/impulsivity). As well as local global biological motion perception is reduced in ADHD participants. In addition, the study demonstrates a significant correlation between local biological motion perception skills and the social responsiveness score in the ADHD group, but not the controls. A path analysis in the ADHD data suggests that general performance in biological motion perception is influenced mainly by global biological motion perception performance and attentional and perceptual reasoning skills.

Strengths:

It is true that there exists not much work on biological motion perception and ADHD. Therefore, the presented study contributes an interesting new result to the biological motion literature and adds potentially also new behavioral markers for this clinical condition. The design of the study is straightforward and technically sound, and the drawn conclusions are supported by the presented results.

Thank you for your positive assessment of our work.

Weaknesses:

Some of the claims about the relationship between genetic factors and ADHD and the components of biological motion processing have to remain speculative at this point because genetic influences were not explicitly tested in this paper.

We agree that the relationship between genetic factors and BM processing in ADHD needs more investigation, We have modified our statement in Discussion section as following:

“Using the classical twin method, Wang et al. found that the distinction between local and global BM processing may stem from the dissociated genetic bases. The former, to a great degree, seems to be acquired phylogenetically20,21,59,60, while the latter is primarily obtained through individual development19.” (lines 421 - 425),

Reviewer #2 (Public Review):

Summary:

Tian et al. aimed to assess differences in biological motion (BM) perception between children with and without ADHD, as well as relationships to indices of social functioning and possible predictors of BM perception (including demographics, reasoning ability and inattention). In their study, children with ADHD showed poorer performance relative to typically developing children in three tasks measuring local, global, and general BM perception. The authors further observed that across the whole sample, performance in all three BM tasks was negatively correlated with scores on the social responsiveness scale (SRS), whereas within groups a significant relationship to SRS scores was only observed in the ADHD group and for the local BM task. Local and global BM perception showed a dissociation in that global BM processing was predicted by age, while local BM perception was not. Finally, general (local & global combined) BM processing was predicted by age and global BM processing, while reasoning ability mediated the effect of inattention on BM processing.

Strengths:

Overall, the manuscript is presented in a relatively clear fashion and methods and materials are presented with sufficient detail so the study could be reproduced by independent researchers. The study uses an innovative, albeit not novel, paradigm to investigate two independent processes underlying BM perception. The results are novel and have the potential to have wide-reaching impact on multiple fields.

We appreciate your positive assessment of our work.

Weaknesses:

Except for the main analysis, it is unclear what the authors' specific predictions are regarding the three different tasks they employ. The three BM tasks are used to probe different processes underlying BM perception, but it is difficult to gather from the introduction why these three specific tasks were chosen and what predictions the authors have about the performance of the ADHD group in these tasks. Relatedly, the authors do not report whether (and if so, how) they corrected for multiple comparisons in their analyses. As the number of tests one should control for depends on the theoretical predictions (http://daniellakens.blogspot.com/2016/02/why-you-dont-need-to-adjust-you-alpha.html), both are necessary for the reader to assess the statistical validity of the results and any inferences drawn from them. The same is the case for the secondary analyses exploring relationships between the 3 individual BM tasks and social function measured by the social responsivity scale (SRS).

We appreciate these constructive suggestions. In response, we have included a detailed description in the Introduction section explaining why we employed three different tasks and our predictions about the performance in ADHD:

“Despite initial indications, a comprehensive investigation into BM perception in ADHD is warranted. We proposed that it is essential to deconstruct BM processing into its multiple components and motion features, since treating them as a single entity may lead to misleading or inconsistent findings31. To address this issue, we employed a carefully designed behavioral paradigm used in our previous study19, making slight adjustments to adapt for children. This paradigm comprises three tasks. Task 1 (BM-local) aimed to assess the ability to process local BM cues. Scrambled BM sequences were displayed and participants could use local BM cues to judge the facing direction of the scrambled walker. Task 2 (BM-global) tested the ability to process the global configuration cues of the BM walker. Local cues were uninformative, and participants used global BM cues to determine the presence of an intact walker. Task 3 (BM-general) tested the ability to process general BM cues (local + global cues). The stimulus sequences consisted of an intact walker and a mask containing similar target local cues, so participants could use general BM cues (local + global cues) to judge the facing direction of the walker.” (lines 116 - 130)

“In Experiment 1, we examined three specific BM perception abilities in children with ADHD. As mentioned earlier, children with ADHD also show impaired social interaction, which implies atypical social cognition. Therefore, we speculated that children with ADHD performed worse in the three tasks compared to TD children.” (lines 131 - 134)

Additionally, we have reported the p values corrected for multiple comparisons (false discovery rate, FDR) in the revised manuscript wherever it was necessary to adjust the alpha (lines 310 - 316; Table 2). The pattern of the results remained unchanged.

In relation to my prior point, the authors could provide more clarity on how the conclusions drawn from the results relate to their predictions. For example, it is unclear what specific conclusions the authors draw based on their findings that ADHD show performance differences in all three BM perception tasks, but only local BM is related to social function within this group. Here, the claim is made that their results support a specific hypothesis, but it is unclear to me what hypothesis they are actually referring to (see line 343 & following). This lack of clarity is aggravated by the fact that throughout the rest of the discussion, in particular when discussing other findings to support their own conclusions, the authors often make no distinction between the two processes of interest. Lastly, some of the authors' conclusions related to their findings on local vs global BM processing are not logically following from the evidence: For instance, the authors conclude that their data supports the idea that social atypicalities are likely to reduce with age in ADHD individuals. However, according to their own account, local BM perception - the only measure that was related to social function in their study - is understood to be age invariant (and was indeed not predicted by age in the present study).

Thank you for pointing out this issue. We have carefully revised the Discussion section about our findings to clarify these points:

“Our study contributes several promising findings concerning atypical biological motion perception in ADHD. Specifically, we observe the atypical local and global BM perception in children with ADHD. Notably, a potential dissociation between the processing of local and global BM information is identified. The ability to process local BM cues appears to be linked to the traits of social interaction among children with ADHD. In contrast, global BM processing exhibits an age-related development. Additionally, general BM perception may be affected by factors including attention.” (lines 387 - 393)

We have provided a detailed discussion on the two processes of interest to clarify their potential differences and the possible reasons behind the difference of the divergent developmental trajectories between local and global BM processing:

“BM perception is considered a multi-level phenomenon56-58. At least in part, processing information of local BM and global BM appears to involve different genetic and neural mechanisms16,19. Using the classical twin method, Wang et al. found that the distinction between local and global BM processing may stem from the dissociated genetic bases. The former, to a great degree, seems to be acquired phylogenetically20,21,59,60, while the latter is primarily obtained through individual development19. The sensitivity to local rather than global BM cues seems to emerge early in life. Visually inexperienced chicks exhibit a spontaneous preference for the BM stimuli of hen, even when the configuration was scrambled20. The same finding was reported in newborns. On the contrary, the ability to process global BM cues rather than local BM cues may be influenced by attention28,29 and shaped by experience24,56.” (lines 419 - 430)

“We found that the ability to process global and general BM cues improved significantly with age in both TD and ADHD groups, which imply the processing module for global BM cues tends to be mature with development. In the ADHD group, the improvement in processing general and global BM cues is greater than that in processing local BM cues, while no difference was found in TD group. This may be due to the relatively higher baseline abilities of BM perception in TD children, resulting in a relatively milder improvement. These findings also suggest a dissociation between the development of local and global BM processing. There seems to be an acquisition of ability to process global BM cues, akin to the potential age-related improvements observed in certain aspects of social cognition deficits among individuals with ADHD5, whereas local BM may be considered an intrinsic trait19.” (lines 438 -449)

In addition, we have rephased some inaccurate statements in revised manuscript. Another part of social dysfunction might be stable and due to the atypical local BM perception in ADHD individuals, although some studies found a part of social dysfunction would reduce with age in ADHD individuals. One reason is that some factors related to social dysfunction would improve with age, like the symptom of hyperactivity.

Results reported are incomplete, making it hard for the reader to comprehensively interpret the findings and assess whether the conclusions drawn are valid. Whenever the authors report negative results (p-values > 0.05), the relevant statistics are not reported, and the data not plotted. In addition, summary statistics (group means) are missing for the main analysis.

Thanks for your comments. We have provided the complete statistical results in the revised manuscript (lines 309 - 316) and supplementary material, which encompass relevant statistics and plots of negative results (Figure 4, Figure S2 and S3), in accordance with our research questions. And we have also included summary statistics in the Results section (lines 287 - 293).

Some of the conclusions/statements in the article are too strong and should be rephrased to indicate hypotheses and speculations rather than facts. For example, in lines 97-99 the authors state that the finding of poor BM performance in TD children in a prior study 'indicated inferior applicability' or 'inapplicable experimental design'. While this is one possibility, a perhaps more plausible interpretation could be that TD children show 'poor' performance due to outstanding maturation of the underlying (global) BM processes (as the authors suggest themselves that BM perception can improve with age). There are several other examples where statements are too strong or misleading, which need attention.

We thank you for pointing out the issue. We have toned down and rephrased the strong statements and made the necessary revisions.

“Another study found that children with ADHD performed worse in BM detection with moderate ratios of noise34. This may be due to the fact that BM stimuli with noise dots will increase the difficulty of identification, which highlights the difference in processing BM between the two groups33,35.” (lines 111 - 115)

Reviewer #3 (Public Review):

Summary:

The authors presented point light displays of human walkers to children (mean = 9 years) with and without ADHD to compare their biological motion perception abilities and relate them to IQ, social responsiveness scale (SRS) scores and age. They report that children with ADHD were worse at all three biological motion tasks, but that those loading more heavily on local processing related to social interaction skills and global processing to age. The important and solid findings are informative for understanding this complex condition, as well as biological motion processing mechanisms in general. However, I am unsure that these differences between local and global skills are truly supported by the data and suggest some further analyses.

Strengths:

The authors present clear differences between the ADHD and TD children in biological motion processing, and this question has not received as much attention as equivalent processing capabilities in autism. They use a task that appears well controlled. They raise some interesting mechanistic possibilities for differences in local and global motion processing, which are distinctions worth exploring. The group differences will therefore be of interest to those studying ADHD, as well as other developmental conditions, and those examining biological motion processing mechanisms in general.

We appreciate your positive feedback. In revised manuscript, we have added more analyses to support the differences between local and global motion processing. Please refer to our response to the point #3 you mentioned below.

Weaknesses:

I am unsure that the data are strong enough to support claims about differences between global and local processing wrt social communication skills and age. The mechanistic possibilities for why these abilities may dissociate in such a way are interesting, but do not seem so plausible to me. I am also concerned about gender, and possible autism, confounds when examining the effect of ADHD. Specifics:

Gender confound. There are proportionally more boys in the ADHD than TD group. The authors appear to attempt to overcome this issue by including gender as a covariate. I am unsure if this addresses the problem. The vast majority of participants in the ADHD group are male, and gender is categorically, not continuously, defined. I'm pretty sure this violates the assumptions of ANCOVA.

We appreciate your comments. We concur with you that although we observed a clear difference between local and global BM processing in ADHD, the evidence is to some extent preliminary. The mechanistic possibilities for why these abilities may dissociate have been discussed in revised manuscript. Please refer to the response to reviewer 2’s point #2. To further examine if gender played a role in the observed results, we used a statistical matching technique to obtain a sub-dataset. The pattern of results remained with the more balanced dataset (see Supplementary Information part 1). According to your suggestion, we have also presented the results without using gender as a covariate in main text and also separated the data of boys and girls on the plots (see Figure 1 and Figure S1). There were indeed no signs of a gender effect.

Autism. Autism and ADHD are highly comorbid. The authors state that the TD children did not have an autism or ADHD diagnosis, but they do not state that the ADHD children did not have an autism diagnosis. Given the nature of the claims, this seems crucial information for the reader.

Thanks for your suggestion. We have confirmed that all children with ADHD in our study were not diagnosed with autism. We used a semi-structured interview instrument (K-SADSPL-C) to confirm every recruited child with ADHD but not with ASD. The exclusion criteria for both groups were mentioned in the Materials and methods section:

“Exclusion criteria for both groups were: (a) neurological diseases; (b) other neurodevelopmental disorders (e.g., ASD, Mental retardation, and tic disorders), affective disorders and schizophrenia…” (lines 158 - 162)

Conclusions. The authors state frequently that it was the local BM task that related to social communication skills (SRS) and not the global tasks. However, the results section shows a correlation between SRS and all three tasks. The only difference is that when looking specifically within the ADHD group, the correlation is only significant for the local task. I think that if the authors wish to make strong claims here they must show inferential stats supporting (1) a difference between ADHD and TD SRS-Task 1 correlations, and (2) a difference in those differences for Task 2 and 3 relative to Task 1. I think they should also show a scatterplot of this correlation, with separate lines of best fit for the two groups, for Tasks 2 and 3 as well. I.e. Figure 4 should have 3 panels. I would recommend the same type of approach for age. Currently, they have small samples for correlations, and are reading much of theoretical significance between some correlations passing significance threshold and others not. It would be incredibly interesting if the social skills (as measured by SRS) only relate to local BM abilities, and age only to global, but I think the data are not so clear with the current information. I would be surprised if all BM abilities did not improve with age. Even if there is some genetic starter kit (and that this differs according to particular BM component), most abilities improve with learning/experience/age.

Thank you for this recommendation. We have added more statistics to test differences between the correlations (a difference between ADHD and TD in SRS-Task 1 correlations (see the first paragraph of Supplementary Information part 2), a difference in SRS-response accuracy correlations for Task 2 and 3 relative to Task 1(see the second paragraph of Supplementary Information part 2), and a difference in age-response accuracy correlations for Task 2 and 3 relative to Task 1 in ADHD group (see Supplementary Information part 3)). Additionally, we have included scatterplots for SRS-Task1, SRS-Task2, SRS-Task3 (with separate lines of best fit for the two groups in each, see Figure 4), SRS-ADHD, SRS-TD, age-ADHD and age-TD (with separate lines of best fit for the three tasks in each, see Figure S2 and S3) to make a clear demonstration. Detailed results have been presented in the revised manuscript and Supplementary Information. We expect these further analyses would strengthen our conclusions.

Theoretical assumptions. The authors make some sweeping statements about local vs global biological motion processing that need to be toned down. They assume that local processing is specifically genetically whereas global processing is a product of experience. The fact their global, but not local, task performance improves with age would tend to suggest there could be some difference here, but the existing literature does not allow for this certainty. The chick studies showing a neonatal preference are controversial and confounded - I cannot remember the specifics but I think there an upper vs lower visual field complexity difference here.

Thank you for pointing out this issue. We have toned down rephrased our claims that the difference between local and global BM processing according to your suggestion:

“These findings suggest that local and global mechanisms might play different roles in BM perception, though the exact mechanisms underlying the distinction remain unclear. Exploring the two components of BM perception will enhance our understanding of the difference between local and global BM processing, shedding light on the psychological processes involved in atypical BM perception.” (lines 87 - 92)

Reviewer #1 (Recommendations For The Authors):

I have only a number of minor points that should be addressed prior to publication:

L. 95ff: What is meant by 'inapplicability of experimental designs' ? This paragraph is somewhat unclear.

In revised manuscript, we have clarified this point (lines 111 - 115).

L. 146: The groups were not perfectly balanced for sex. Would results change fundamentally in a more balanced design, or can arguments be given that gender does not play a role, like it seems to be the case for some functions in biological motion perception (e.g. Pavlova et al. 2015; Tsang et al 2018). One could provide a justification that this disbalance does not matter or test for subsampled balanced data sets maybe.

This point is similar to the point #1 from reviewer 3, and we have addressed this issue in our response above.

L. 216 f.: In this paragraph it does not become very clear that the mask for the global task consisted of scrambles generated from walkers walking in the same direction. The mask for the local task then should consist of a balanced mask that contains the same amount of local motion cues indicating right and leftwards motion. Was this the case? (Not so clear from this paragraph.)

Regarding the local task, the introduction of mask would make the task too difficult for children. Therefore, in the local task, we only displayed a scrambled walker without a mask, which was more suitable for children to complete the task. We have made clear this point in the corresponding paragraph (lines 232 - 241).

L. 224 ff.: Here it would be helpful to see the 5 different 'facing' directions of the walkers. What does this exactly mean? Do they move on oblique paths that are not exactly orthogonal to the viewing directions, and how much did these facing directions differ?

Out of the five walkers we used, two faced straight left or right, orthogonal to the viewing directions. Two walked with their bodies oriented 45 degrees from the observer, to the left or right. The last one walked towards the observer. We have included a video (Video 4) to demonstrate the 5 facing directions.

L. 232: How was the number of 5 practicing trials determined/justified?

As mentioned in main text, global BM processing is susceptible to learning. Therefore, too many practicing trials would increase BM visual experience and influence the results. We determined the number of training trials to be 5 based on the results of the pilot experiment. During this phase, we observed that nearly all children were able to understand the task requirements well after completing 5 practicing trials.

L 239: Apparently no non-parametric statistics was applied. Maybe it would be good to mention in the Statistics section briefly why this was justified.

We appreciate your suggestion and have cited two references in the Statistics section (Fagerland et al. 2012, Rochon et al. 2012). Fagerland et al., mentioned that when the sample size increases, the t-test is more robust. According to the central limit theorem, when the sample size is greater than 30, the sampling distribution of the mean can be safely assumed to be normal.

(http://www2.psychology.uiowa.edu/faculty/mordkoff/GradStats/part%201/I.07%20normal.p df). In fact, we also ran non-parametric statistics for our data and found the results to be robust.

L 290: 'FIQ' this abbreviation should be defined.

Regarding the abbreviation ’FIQ’, it stands for the abbreviation of the full-scale intellectual quotient, which was mentioned in Materials and methods section:

“Scores of the four broad areas constitute the full-scale intellectual quotient (FIQ).”

L. 290 ff.: These model 'BM-local = age + gender etc ' is a pretty sloppy notation. I think what is meant that a GLM was used that uses the predictors gender etc. time appropriate beta_i values. This formula should be corrected or one just says that a GLM was run with the predictors gender ....

The same criticism applies to these other models that follow.

We thank you for pointing this out. We have modified all formulas accordingly in the revised manuscript (see part3 of the Results section).

All these models assume linearity of the combination of the predictors.was this assumption verified?

We referred to the previous study of BM perception in children. They found main predictor variables, including IQ (Rutherford et al., 2012; Jones et al., 2011) and age (Annaz et al., 2010; van et al., 2016), have a linear relation with the ability of BM processing.

L. 296ff.: For model (b) it looks like general BM performance is strongly driven by the predictor global BM performance in the group of patients. Does the same observation also apply to the normals?

The same phenomenon was not observed in TD children. We have briefly discussed this point in the Discussion section of the revised manuscript (lines 449 - 459).

Reviewer #2 (Recommendations For The Authors):

(1) Please add public access to the data repository so data availability can be assessed.

The data of the study will be available at https://osf.io/37p5s/.

(2) Although overall, the language was clear and understandable, there are a few parts where language might confuse a reader and lead to misconceptions. For instance, line 52: Did the authors mean to refer to 'emotions and intentions' instead of 'emotions and purposes'? See also examples where rephrasing may help to reflect a statement is speculation rather than fact.

Thanks for the comments. We have carefully checked the full text and rephrased the confused statements.

(3) Line 83/84: Autism is not a 'mental disorder' - please change to something like 'developmental disability'. Authors are encouraged to adapt their language according to terms preferred by the community (e.g., see Fig. 5 in this article:

https://onlinelibrary.wiley.com/doi/10.1002/aur.2864)

Suggestion well taken. We have changed the wording accordingly:

“In recent years, BM perception has received significant attention in studies of mental disorders (e.g., schizophrenia30) and developmental disabilities, particularly in ASD, characterized by deficits in social communication and social interaction31,32.” (lines 93 - 95)

(4) Please report how the sample size for the study was determined.

In the Materials and methods section (lines 168 - 173), we explained how the sample size was determined.

Line 94: It would be helpful to have a brief description of what neurophysiological differences have been observed upon BM perception in children with ADHD.

Thanks for the comment. We have added a brief description of neurophysiological findings in children with ADHD (lines 108 - 111).

(6) Line 106/107 and 108/109: please add references.

We have revised this part, and the relevant findings and references are in line with the revised manuscript (lines 77, 132 - 133).

(7) Line 292: Please add what order the factors were entered into each regression model.

Regarding this issue, we used SPSS 26 for the main analysis. SPSS utilizes the Type III sum of squares (default) to evaluate models. Regardless of the order in the GLM, we will obtain the same result. For more information, please refer to the documentation of SPSS 26 (https://www.ibm.com/docs/en/spss-statistics/26.0.0?topic=features-glm-univariate-analysis).

Reviewer #3 (Recommendations For The Authors)

(1) Task specifics. It is key to understanding the findings, as well as the dissociation between tasks, that the precise nature of the stimuli is clear. I think there is room for improvement in description here. Task 1 is described as involving relocating dots within the range of the intact walker. Of course, PLWs are created by presenting dots at the joints, so relocation can involve either moving to another place on the body, or random movement within the 2D spatial array (which likely involves moving it off the body). Which was done? It is said that Ps must indicate the motion direction, but what was the display of the walker? Sagittal? Task 2 requires detecting whether there is an intact walker amongst scrambled walkers. Were all walkers completely overlaid? Task 3 requires detecting the left v right facing of an intact walker at different orientations, presented amongst noise. So Task 3 requires determining facing direction and Task 1 walking direction. Are these tasks the same but described differently? Or can walkers ever walk backwards? Wrt this point, I also think it would help the reader if example videos were uploaded.

We appreciate you for bringing this to our attention. With regards to Task 1, it appears that your second speculation is correct. We scrambled the original dots and randomly presented them within the 2D spatial array (which likely involved moving them off the body). As a result, the global configuration of the 13 dots was completed disrupted while preserving the motion trajectory of each individual dot. This led to the display of scrambled dots on the monitor (which does not resemble a human). In practice, these local BM cues contain information about motion direction. In Task 2, the target walkers completely overlaid by a mask that is approximately 1.44 times the size of the intact walker. The task requirements of Task 1 and Task3 are same, which is judging the motion (walking) direction. The difference is that Task 1 displayed a scrambled walker while Task 3 displayed an intact walker within a mask. We have clarified these points and improved our descriptions in Procedure section and created example videos for each task, which we believe will be helpful for the readers to understand each task.

(2) Gender confound (see above). I think that the authors should present the results without gender as a covariate. Can they separate boys and girls on the plots with different coloured individual datapoints, such that readers can see whether it's actually a gender effect driving the supposed ADHD effect? And show that there are no signs of a gender effect in their TD group?

This point is similar to the point #1 you mentioned. Please refer to our response to that point above.

(3) Autism possible confound (see above). I think the authors must report whether any of the ADHD group had an autism diagnosis.

Please refer to the response for the point #2 your mentioned.

(4) Conclusions concerning differences between the local and global tasks wrt SRS and age (see above). I believe the authors should add stats demonstrating differences between the correlations to support such claims, as well as demonstrating appropriate scatterplots for SRS-Task 1, SRS-Task 2, SRS-Task 3 and age-Task 1, age-Task2 and age-Task 3 (with separate lines of best fit for the two groups in each).

Please refer to the response for the point #3 your mentioned.

(5) Theoretical assumptions (see above). I would suggest rephrasing all claims here to outline that these discussed mechanistic differences between local and global BM processing are only possibilities and not known on the basis of existing data.

Please refer to the response for the point #4 your mentioned.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The ability to perceive biological motion (BM) is crucial for human survival and social interaction. Plentiful studies have found impaired BM perception in autism spectrum disorder characterized by deficits in social interaction. Children with attention deficit hyperactivity disorder (ADHD) often exhibit similar deficits in social interaction, but few studies have investigated BM perception in ADHD. Here, we compared the differences in abilities to process local kinematic and global configurational cues, two fundamental abilities of BM perception, between typically developing (TD) and ADHD children. Then, we further investigated the relationship between BM perception and social interaction skills measured by the Social Responsiveness Scale (SRS) and examined the contributions of potential factors (e.g., gender, age, attention and intelligence) to BM perception. Results revealed that children with ADHD exhibited atypical BM perception with a clear dissociation between local and global BM information processing. Local BM processing ability was largely related to social interaction skills, whereas global BM processing ability would significantly improve with age. Critically, general BM perception (i.e., both local and global BM cues) was affected by sustained attention ability in ADHD. This relationship was mainly mediated by Reasoning Intelligence. These findings elucidate the atypical biological motion perception in ADHD and the potential factors related to BM perception. Moreover, this study provides new evidence for local BM perception as a hallmark of social cognition and advances the comprehensive understanding of the distinct roles of local and global processing in BM perception and social cognitive disorders.

Revised version: The ability to perceive biological motion (BM) is crucial for human survival and social interaction. Plentiful studies have found impaired BM perception in autism spectrum disorder characterized by deficits in social interaction. Children with attention deficit hyperactivity disorder (ADHD) often exhibit similar deficits in social interaction, but few studies have investigated BM perception in ADHD. Here, we compared the differences in abilities to process local kinematic and global configurational cues, two fundamental abilities of BM perception, between typically developing (TD) and ADHD children. Then, we further investigated the relationship between BM perception and social interaction skills measured by the Social Responsiveness Scale (SRS) and examined the contributions of latent factors (e.g., gender, age, attention and intelligence) to BM perception. Results revealed that children with ADHD exhibited atypical BM perception with a potential dissociation between local and global BM information processing. Local BM processing ability was largely related to social interaction skills, whereas global BM processing ability would significantly improve with age. Critically, general BM perception (i.e., both local and global BM cues) could be affected by sustained attention ability in children with ADHD. This relationship was mainly mediated by Reasoning Intelligence. These findings elucidate the atypical biological motion perception in ADHD and the latent factors related to BM perception. Moreover, this study provides new evidence for local BM perception as a hallmark of social cognition and advances the comprehensive understanding of the distinct roles of local and global processing in BM perception and social cognitive disorders.

Introduction

Earlier version: Attention deficit/hyperactivity disorder (ADHD) is a common developmental disorder. The prevalence of ADHD in children and adolescents is between 2% and 7%, with an average of around 5%1. Besides the well-established core symptoms of ADHD (including a deficit of sustained attention, hyperactivity, and impulsivity), some ASD characteristics, such as dysfunctions in social communication and social interaction, have also been frequently reported in ADHD children2–4. However, experimental studies of social cognition in ADHD are still scarce. Some studies described worse performance in tasks of social cognition. Among them, impaired theory of mind (ToM) and emotion recognition are most frequently reported5–7. It is difficult for children with ADHD to recognize others’ emotions and purposes. Other social cognition processes in ADHD are unclear. Further exploring diversified social cognition (e.g., biological motion perception) will provide a new understanding of impaired social function in ADHD. Moreover, recent studies revealed that the social cognition in ADHD can potentially vary with different factors at the cognitive, pathological, or developmental levels, such as general cognitive impairment5, the severity of symptoms8, or age5. Yet, how these factors relate to the dysfunctions of social cognition in ADHD is still in its infancy. Bridging the gap is important because this knowledge can help depict the developmental trajectory of social cognition and find effective interventions for impaired social interaction in ADHD.

Revised version: Attention deficit/hyperactivity disorder (ADHD) is a common developmental disorder, with a prevalence ranging from 2% to 7% in children and adolescents, averaging around 5%1. Besides the well-established core symptoms of ADHD (including a deficit of sustained attention, hyperactivity, and impulsivity), some autism spectrum disorder (ASD) characteristics, such as dysfunctions in social communication and social interaction, have also been frequently reported in ADHD children2–4. However, experimental studies of social cognition in ADHD are still scarce. Some studies described worse performance in tasks of social cognition. Among them, impaired theory of mind (ToM) and emotion recognition are most frequently reported5–7. It is difficult for children with ADHD to recognize others’ emotions and intentions. However, the understanding of other social cognition processes in ADHD remains limited. Further exploring a diverse range of social cognition (e.g., biological motion perception) can provide a fresh perspective on the impaired social function observed in ADHD. Moreover, recent studies indicated that the social cognition in ADHD may vary depending on different factors at the cognitive, pathological, or developmental levels, such as general cognitive impairment5, the severity of symptoms8, or age5. Yet, how these factors relate to the dysfunctions of social cognition in ADHD is still in its infancy. Bridging the gap is important as it can help depict the developmental trajectory of social cognition and identify effective interventions for impaired social interaction in individuals with ADHD.

Earlier version: Biological motion (BM), namely the movement of a living creature, carries much information beyond bodily movements9, such as intention10, emotion11, gender12, and identity13, 14. Since point-light displays (PLD) technology was used to depict human motions15, it is possible to separate biological motion from other characteristics like shape and color. Considering the seminal impact on cognitive science, developmental and clinical neuroscience, BM perception has drawn significant attention from scientists. Some researchers attempted to deconstruct BM processing into more specific components. Our study concentrated on two fundamental abilities to process BM cues (Figure 1): the abilities to process local BM cues derived from major joint motion tracks and global BM cues of human configuration. Separable neural signals revealed two independent processing stages16–18. The ability to process local BM cues is genetically highly preserved19; for example, newly hatched chicks without visual experience exhibit a spontaneous preference for BM20. A similar finding is also reported in 2-day-old babies21. Local BM cues not only carry locomotive direction22 but also aid in detecting life in the visual environment23 without observers’ explicit recognition or attention24–27. Moreover, the processing of local BM cues is unaffected by attention, relatively robust to masking noise and does not show a learning trend27, 28. In contrast, global BM processing involves top-down processing, and attention plays a critical role in global BM perception28, 29. Dispersed attention adversely affects its performance. Compared with local BM processing, global BM processing is susceptible to learning and is heavily hindered by increasing mask densities28. In recent years, BM perception has received wide attention in studies of mental disorders30, particularly in ASD, characterized by deficits in social communication and social interaction31, 32. An important reason is that BM perception is considered a hallmark of social cognition. Individuals with deficits in BM processing exhibit worse social perception in daily life10. Especially in local BM processing, it was correlated with autistic traits, especially in the subdimension of social communication19. Therefore, it is essential to focus on revealing differences in BM processing between typical and atypical development, which will enhance our comprehension of social dysfunction in atypical development.

Revised version: Biological motion (BM), which refers to the movements of a living creature, conveys a wealth of information beyond bodily movements9, such as intention10, emotion11, gender12, and identity13,14. The advent of point-light displays (PLD) technology used to depict human motions15 allowed researchers to separate biological motion from other characteristics like shape and color. Considering the seminal impact on cognitive science, developmental and clinical neuroscience, BM perception has drawn significant attention from scientists. Some researchers attempted to deconstruct BM processing into more specific components. Our study concentrated on two fundamental abilities involved in processing BM cues (Figure 1): the ability to process local BM cues derived from major joint motion tracks and the ability to process global BM cues of human configuration. Previous studies provided some findings about the difference between local and global BM perception. Separable neural signals for these abilities implied two independent BM processing stages16–18. The ability to process local BM cues appears to be genetically highly preserved19. For example, newly hatched chicks without visual experience exhibit a spontaneous preference for BM20. A similar finding is also reported in 2-day-old babies21. Local BM cues not only carry locomotive direction22 but also contribute to detecting life in the visual environment23 without observers’ explicit recognition or attention24–27. As a result, the processing of local BM cues is unaffected by attention, relatively robust to masking noise and does not show a learning trend27,28. In contrast, global BM processing involves top-down modulation, with attention playing a critical role in its perception28,29. Dispersed attention adversely affects its performance. Compared with local BM processing, global BM processing is susceptible to learning and is heavily hindered by increasing mask densities28. These findings suggest that local and global mechanisms might play different roles in BM perception, though the exact mechanisms underlying the distinction remain unclear. Exploring the two components of BM perception will enhance our understanding of the difference between local and global BM processing, shedding light on the psychological processes involved in atypical BM perception.

Removed in the revised version: Schematic representation of Biological Motion (BM) and scrambled BM sequence.

Earlier version: Intact walker contains information of local kinematics and global configuration. Local kinematics refers to the motion track of each critical joint illustrated by the chromatic dot. Global configuration is composed of the relative locations of each joint. In scrambled BM sequence, global configuration cues have been deprived, but local kinematics have been retained. (Figure reconstructed from Wang et al., 201819)

Revised version: Schematic representation of Biological Motion (BM) and scrambled BM sequence. Intact walker contains information of local kinematics and global configuration. Local kinematics refers to the motion tracks of each critical joint illustrated by the chromatic dot. Global configuration is composed of the relative locations of each joint. In scrambled BM sequence, global configuration cues have been deprived, but local kinematics have been retained. (Figure reconstructed from Wang et al., 201819)

Added in the revised version: In recent years, BM perception has received significant attention in studies of mental disorders (e.g., schizophrenia30) and developmental disabilities, particularly in ASD, characterized by deficits in social communication and social interaction31,32. An important reason is that BM perception is considered a hallmark of social cognition. Individuals with deficits in BM processing exhibit worse social perception in daily life10. The other study found that participants’ ability to process BM cues was correlated with their autistic traits, especially in the subdimension of social communication19. Notably, the contributing factors of impaired BM perception in children with ASD is unclear. Previous meta-analysis studies of BM perception in ASD found mixed results about the effects of age and IQ on processing BM31,32.

Added in the revised version: Therefore, it is essential to focus on revealing differences in BM processing between typical and atypical development and identified potential factors that may influence BM perception, which will enhance our comprehension of social dysfunction in atypical development. Compared with numerous studies examining impaired BM perception in ASD, few studies focused on BM perception in children with ADHD. An EEG study found neuro-electrophysiological changes in processing BM stimuli in ADHD33. Specifically, children with ADHD showed reduced activity in motion- sensitive components (N200) while watching biological and scrambled motion, although no behavioral difference was observed. Another study found that children with ADHD performed worse in BM detection with moderate ratios of noise34. This may be due to the fact that BM stimuli with noise dots will increase the difficulty of identification, which highlights the difference in processing BM between the two groups33,35.

Added in the revised version: Despite initial indications, a comprehensive investigation into BM perception in ADHD is warranted. We proposed that it is essential to deconstruct BM processing into its multiple components and motion features, since treating them as a single entity may lead to misleading or inconsistent findings31. To address this issue, we employed a carefully-designed behavioral paradigm used in our previous study19, making slight adjustments to adapt for children. This paradigm comprises three tasks. Task 1 (BM-local) aimed to assess the ability to process local BM cues. Scrambled BM sequences were displayed and participants could use local BM cues to judge the facing direction of the scrambled walker. Task 2 (BM-global) tested the ability to process the global configuration cues of the BM walker. Local cues were uninformative, and participants used global BM cues to determine the presence of an intact walker. Task 3 (BM-general) tested the ability to process general BM cues (local + global cues). The stimulus sequences consisted of an intact walker and a mask containing similar target local cues, so participants could use general BM cues (local + global cues) to judge the facing direction of the walker.

Added in the revised version: In Experiment 1, we examined three specific BM perception abilities in children with ADHD. As mentioned earlier, children with ADHD also show impaired social interaction2–4, which implies atypical social cognition. Therefore, we speculated that children with ADHD performed worse in the three tasks compared to TD children. In Experiment 2, we further explored the relationship between BM perception and social interaction ability in ADHD, as well as identified potential factors (e.g., IQ, age and attention) that may affect BM perception in this population. We speculate that if local motion and global configuration of BM are indeed contributed by distinct mechanisms as suggested by previous studies, then their impairments in the ADHD population and the influential factors behind the impairments may be different.

Removed in the revised version: Few studies focused on BM perception in children with ADHD, compared with numerous studies of impaired BM perception in ASD. An EEG study found neuro-electrophysiological changes in processing BM stimuli in ADHD33, whereas no behavioral difference existed, which could be improved by more carefully designed BM stimuli. Another study found that children with ADHD performed worse in BM detection with moderate ratios of noise34. Nevertheless, the accuracy of TD groups was also very low (i.e., poor performance), which indicated the inferior applicability of this paradigm. In short, factors such as inapplicable experimental designs may influence the precision and reliability of the results. Despite some preliminary hints, systematic evaluation of BM perception in ADHD is warranted.

Removed in the revised version: Therefore, the current study aims to examine BM perception abilities in children with ADHD. The first question to be addressed by this study is whether the BM perception abilities differ between TD and ADHD children. As mentioned above, BM perception is related to social interaction and has two fundamental components. Local BM processing ability is genetically highly preserved, while global BM processing ability shows a learning trend and needs the involvement of attention. ADHD is a neurodevelopmental disorder characterized by attention deficit and mostly accompanied by impaired social interaction. Therefore, the second question is whether BM perception is associated with social interaction ability in ADHD, and if so, which factors (e.g., attention) could affect BM perception in ADHD. To answer these questions, we utilized a stable and ingenious behavioral paradigm used in a previous study and made minor modifications to fit it to children19. This paradigm combined an intact and a scrambled point-lighted walker with a noise mask to measure local and global BM information processing abilities. Using this paradigm, in Experiment 1, we compared the performance in processing two different BM cues (i.e., local and global BM cues) as well as their combination (i.e., both local and global cues) between TD and ADHD groups. Then, in Experiment 2, we collected the social interaction trait of participants to examine its correlation with the performance of BM tasks. Finally, we investigated the contributing factors of BM perception in ADHD and explored their potential relations.

Materials and methods

Participants

Earlier version: We recruited a total of one hundred and seventeen children with and without ADHD to enter this study. Eighty-one children met the ADHD diagnostic criteria of the Diagnostic and Statistical Manual of Mental Disorders (DSM-5)35. The clinical diagnosis was first made by an experienced child and adolescent psychiatrist in the Child and Adolescent Psychiatric Outpatient Department of the Peking University Sixth Hospital, based on the ADHD Rating Scale. The Chinese version of the kiddie-schedule for affective disorders and schizophrenia-present and lifetime version DSM-5 (K-SADS-PL-C DSM-5)36, 37, a semi-structured interview instrument, was then implemented to confirm the diagnosis. Thirty-six typically developing (TD) children were from ordinary primary schools in Beijing and screened out the presence of ADHD, ASD, affective disorders and behavioral disorders by a trained psychiatrist. All subjects of ADHD groups have full-scale IQ above 75 (5th upper percentile) on Wechsler Intelligence Scale for Children-Fourth Edition, and all TD children had a full-scale IQ that fell above the 5th percentile on the Raven’s Standard Progressive Matrices38 which is used in measuring reasoning ability and regarded as a non-verbal estimate of intelligence. Exclusion criteria for both groups included: (a) neurological diseases; (b) other neurodevelopmental disorders (e.g., ASD, Mental retardation, and tic disorders), affective disorders and schizophrenia; (c) that would impact experiments completion; (d) taking psychotropic drugs or stimulants in 30 days. (e) previous head trauma or neurosurgery.

Revised version: We recruited a total of one hundred and seventeen children with and without ADHD to enter this study. Eighty-one children met the ADHD diagnostic criteria of the Diagnostic and Statistical Manual of Mental Disorders (DSM-5)36. The clinical diagnosis was first made by an experienced child and adolescent psychiatrist in the Child and Adolescent Psychiatric Outpatient Department of the Peking University Sixth Hospital, based on the ADHD Rating Scale. The Chinese version of the kiddie- schedule for affective disorders and schizophrenia-present and lifetime version DSM- 5 (K-SADS-PL-C DSM-5)37,38, a semi-structured interview instrument, was then implemented to confirm the diagnosis. Thirty-six typically developing (TD) children were from ordinary primary schools in Beijing and screened out the presence of ADHD, ASD, affective disorders and behavioral disorders by a trained psychiatrist. All subjects of ADHD groups have full-scale IQ above 75 (5th upper percentile) on Wechsler Intelligence Scale for Children-Fourth Edition, and all TD children had a full-scale IQ that fell above the 5th percentile on the Raven’s Standard Progressive Matrices39 which is used in measuring reasoning ability and regarded as a non-verbal estimate of intelligence. Exclusion criteria for both groups were: (a) neurological diseases; (b) other neurodevelopmental disorders (e.g., ASD, Mental retardation, and tic disorders), affective disorders and schizophrenia; (c) that would impact experiments completion; (d) taking psychotropic drugs or stimulants in 30 days. (e) previous head trauma or neurosurgery.

Earlier version: 36 TD children (age = 9.09 ± 2.18, 14 male) and 39 children with ADHD (age = 9.88 ± 2.23, 28 male) participated in Experiment 1. Groups did not differ by age (t = –1.550, p = .126) but differed in sex (χ2 = 8.964, p = 0.004). 42 ADHD children (age = 9.34 ± 1.89, 27 male) participated in Experiment 2. They were naïve to the task and did not participate in Experiment 1. Participant demographic characteristics are shown in Table 1. All individuals of each group had a normal or corrected-to-normal vision and were naïve to the objectives of experiments. Written informed consent was obtained from the parents of all children before testing. The Institutional Review Board of the Peking University Sixth Hospital has approved this study.

Revised version: Thirty-six TD children (age = 9.09 ± 2.18, 14 male) and thirty-nine children with ADHD (age = 9.88 ± 2.23, 28 male) participated in Experiment 1. Groups did not differ by age (t = -1.550, p = 0.126) but differed in sex (χ2 = 8.964, p = 0.004). 42 ADHD children (age = 9.34 ± 1.89, 27 male) participated in Experiment 2. They were naïve to the task and did not participate in Experiment 1. Participant demographic characteristics are shown in Table 1. There is currently no comparable study in ADHD that indicates the effect size for reference. Studies investigating BM perception in ASD typically have sample sizes ranged from 15 to 35 participants per group32. Considering the mild impairment of social functions in children with ADHD, we determined that a sample size of 35 - 40 participants per group would be reasonable for this study. All individuals of each group had a normal or corrected-to- normal vision and were naïve to the objectives of experiments. Written informed consent was obtained from the parents of all children before testing. The institutional review boards of the Peking University Sixth Hospital and the Institute of Psychology, Chinese Academy of Sciences have approved this study.

Assessment

K-SADS-PL-C DSM-5

Earlier version: K-SADS-PL-C DSM-5 is a semi-structured interview instrument for evaluating mental disorders in children and adolescents aged 6–18 years36. It involved thirty-five diagnoses based on the diagnostic criteria of DSM-5. A trained psychiatrist concluded the diagnosis by interviewing the parents and child. Chinese version showed great psychometric criterion37.

Revised version: K-SADS-PL-C DSM-5 is a semi-structured interview instrument for evaluating mental disorders in children and adolescents aged 6–18 years37. It involved thirty-five diagnoses based on the diagnostic criteria of DSM-5. A trained psychiatrist concluded the diagnosis by interviewing the parents and child. Chinese version showed great psychometric criterion38.

ADHD Rating Scale

Earlier version: ADHD Rating Scale is adapted from the ADHD diagnostic criteria of DSM39, which requires parents or teachers to complete the scale independently. Its Chinese version has an excellent psychometric criterion and consists of 2 subscales40: inattention (IA, nine items) and hyperactivity-impulsivity (HI, nine items). Each item is rated on a four-point Likert scale ranging from 1 (the symptom is “never or rarely”) to 4 (the symptom is “very often”). The final results will create three scores: (1) IA dimension score, (2) HI dimension score, and (3) Total score. Higher scores indicate more severe ADHD symptoms.

Revised version: ADHD Rating Scale is adapted from the ADHD diagnostic criteria of DSM40, which requires parents or teachers to complete the scale independently. Its Chinese version has an excellent psychometric criterion and consists of 2 subscales41: inattention (IA, nine items) and hyperactivity-impulsivity (HI, nine items). Each item is rated on a four-point Likert scale ranging from 1 (the symptom is “never or rarely”) to 4 (the symptom is “very often”). The final results will create three scores: (1) IA dimension score, (2) HI dimension score, and (3) Total score. Higher scores indicate more severe ADHD symptoms.

Social Responsiveness Scale

Earlier version: Social Responsiveness Scale (SRS) is a widely used quantitative measure with 65 items to assess the severity of social impairment in many mental disorders41, and the psychometric properties of the Chinese version are desirable42. It includes five subdimensions: social awareness, social cognition, social communication, social motivation, and autistic mannerisms. Each item is rated on a scale from 0 (never true) to 4 (almost always true), with higher scores indicating worse social ability.

Revised version: Social Responsiveness Scale (SRS) is a widely used quantitative measure with 65 items to assess the severity of social impairment in many mental disorders42, and the psychometric properties of the Chinese version are desirable43. It includes five subdimensions: social awareness, social cognition, social communication, social motivation, and autistic mannerisms. Each item is rated on a scale from 0 (never true) to 3 (almost always true), with higher scores indicating worse social ability.

Wechsler Intelligence Scale for Children-Fourth Edition

Wechsler Intelligence Scale for Children-Fourth Edition (WISC-IV) is widely used to test comprehensive intelligence in individuals aged 6–18. It contains fifteen subtests which comprise four broad areas of intellectual functioning: Verbal Comprehension, Perceptual Reasoning, Working Memory, and Processing Speed. Scores of the four broad areas constitute the full-scale intellectual quotient (FIQ).

QB Test

Earlier version: QB Test is a 15-min Continuous Performance Test (CPT) for assessing inattention and impulsivity with a high-resolution infrared camera monitoring participant’s activity43. Previous psychometric studies have validated its good measurement properties44. After the test was completed, several Q scores were calculated to summarize participant performance. The Q scores were standardized based on normative data matching gender and age. The higher Q score implies more abnormal performance. In this study, we focused on the QbInattention, the Q score responsible for sustained attention, particularly when children focused on tasks.

Revised version: QB Test is a 15-min Continuous Performance Test (CPT) for assessing inattention and impulsivity with a high-resolution infrared camera monitoring participant’s activity44. Previous psychometric studies have validated its good measurement properties45. After the test was completed, several Q scores were calculated to summarize participant performance. The Q scores were standardized based on normative data matching gender and age. The higher Q score implies more abnormal performance. In this study, we focused on the QbInattention, the Q score responsible for sustained attention, particularly when children focused on tasks.

Added in the revised version: Stimuli and Procedure

Removed in the revised version: Stimuli

Earlier version: Stimuli were presented on a 14-inch monitor using MATLAB together with the PsychToolbox extensions. All subjects completed experiments in a dim-lit room with heads on the chinrest to ensure their eyes were 50 cm away from the monitor. Point-light BM stimuli sequences adopted in this study have been used in previous studies45, which were derived from configurations of individual walking motion on a treadmill and did not contain an overall translation. Each frame of BM sequences consisted of 13 white dots representing the human head and major joints and was displayed on a grey background. Each walking cycle lasted 1s with 30 frames. For each trial, the initial frame of BM sequences was randomized. The whole point-lighted BM walker presented about 5.7 ° visual angle vertically.

Revised version: Point-light BM stimuli sequences adopted in this study have been used in previous studies46, which were derived from configurations of individual walking motion on a treadmill and did not contain an overall translation. Each frame of BM sequences consisted of 13 white dots representing the human head and major joints and was displayed on a grey background (see video 1). Each walking cycle lasted 1s with 30 frames. For each trial, the initial frame of BM sequences was randomized. The whole point-lighted BM walker presented about 5.7 ° visual angle vertically. Stimuli were presented on a 14-inch monitor using MATLAB together with the PsychToolbox extensions. All subjects completed experiments in a dim-lit room with heads on the chinrest to ensure their eyes were 50 cm away from the monitor.

Added in the revised version: In Experiment 1, children were required to complete three tasks that were similar to but slightly modified from the versions implemented in our previous study19 (Figure 2). Each trial began with a fixation cross (0.6° * 0.6°). Following a random interval of 1200-1500 ms, the monitor displayed a task-specific BM sequence that lasted 2 seconds (60 frames).

Removed in the revised version: Procedure

Removed in the revised version: In Experiment 1, 36 TD children and 39 children with ADHD were required to complete three tasks that were similar to but slightly modified from the versions implemented in our previous study19 (Figure 2). Each trial began with a fixation cross (0.6° * 0.6°). After a random 1200-1500 ms interval, the monitor displayed a BM sequence specific to each task, lasting 2 s (60 frames). Task 1 (BM-Local) assessed the ability to process local BM cues. 13 dots were randomly relocated within the range of intact BM walker, which rendered global configuration destroyed and local kinematics retained instead. After the display, we required children to press a left or right button to indicate the motion direction of the unidentified creature (i.e., scrambled BM walker) as accurately as possible. Children received no feedback on whether a right or wrong answer was given. There were 30 trials in total and 15 trials for each condition (left and right). Task 2 (BM-Global) tested the ability to process the global configuration cues of the BM walker. Children needed to judge whether there was an intact walker in a mask. Each trial contained a mask consisting of two scrambled target walkers (26 dots) within the boundary about 1.44 times the intact walker. And the scrambled or intact version of the target walker was randomly embedded into the mask. This procedure extracted global configuration cues from the original BM sequence because two conditions (i.e., scrambled and intact walker) contained the same information of local kinematics as the mask, and the local motion cues were hence uninformative. The correct choice asked for using the global component. To avoid learning the shape of the walker24, we set target walkers possibly facing one of five equally spaced directions from left to right. There were 30 trials composed of two conditions (intact or scrambled target). In Task 3 (BM-General), the intact target walker facing left or right was embedded into the mask. The mask is similar to Task 2. Children were required to judge the facing direction of the target walker (left or right). Because the mask was set to covey the local BM cues consistent with the target walker, children could perform the task using general (i.e., local + global) BM information. Task 3 was composed of 30 trials (15 for each facing directions). Other parameters of Task 2 and 3 were similar to Task 1. Before each task, children practiced five trials to make sure good understanding. We performed three tasks in a fixed order so that participants were naïve to the scrambled walker and only made choices using local BM cues in Task 1.

Removed in the revised version: Illustration of the trial sequence.

Earlier version: In Task 1 (BM-local), a monitor displayed scrambled BM sequences. Participants only judged the facing direction of the scrambled walker using local BM cues. In Task 2 (BM-global), each trial only showed an intact or scrambled walker (black dots in the figure) embedded in a mask containing local BM cues. Because the two conditions contained the same local cues that were also present in the mask, a participant must rely on global BM cues to judge whether there was an intact walker in the mask. The figure showed one of five possible facing directions of an intact walker (i.e., facing participants). In Task 3 (BM-general), the stimuli sequence consisted of an intact walker (black dots) and a mask containing similar target local cues, and children judged the facing direction of the walker using general BM cues (local + global). Dots in the figure were rendered in black for better illustration and displayed in white in the actual experiments.

Revised version: Illustration of the trial sequence. In Task 1 (BM-local), a monitor displayed scrambled BM sequences. Participants only judged the facing direction of the scrambled walker using local BM cues. In Task 2 (BM-global), each trial only showed an intact or scrambled walker (black dots in the figure) embedded within a mask containing local BM cues. Because the two conditions contained the same local cues that were also present in the mask, a participant must rely on global BM cues to determine whether an intact walker was present in the mask. The figure showed one of five possible facing directions of an intact walker (i.e., facing participants). In Task 3 (BM-general), the stimuli sequence consisted of an intact walker (black dots) and a mask containing similar target local cues, and children judged the facing direction of the walker using general BM cues (local + global). Dots in the figure were rendered in black for better illustration and displayed in white in the actual experiments.

Added in the revised version: Task 1 (BM-Local) aimed to assess the participants’ ability to process local BM cues. During the task, the monitor only displayed a scrambled walker facing either left or right (see video 2). Specifically, 13 dots constituting the intact walker were randomly relocated within the original range of the BM walker (presented randomly within the 2D region). This manipulation disrupted the global configuration of the intact walker while retaining the local kinematics. After the display, we required children to press a left or right button to indicate the motion direction of the unidentified creature (i.e., scrambled BM walker) as accurately as possible. Children did not receive feedback regarding the accuracy of each response. A total of 30 trials were conducted, with 15 trials for each condition (left and right).

Added in the revised version: Task 2 (BM-Global) tested the ability to process the global configuration cues of the BM walker. A target walker (either scrambled or intact) was displayed within a mask (see video 3) during this task. The mask consisted of two scrambled target walkers (26 dots) with the same locomotion direction as the target walker, displayed within the boundary about 1.44 times larger than the intact walker. The scrambled or intact version of the target walker was randomly embedded into the mask and entirely overlaid by the mask. By this means, the global BM component could be isolated as two conditions (i.e., scrambled and intact walker) contained the same information of local kinematics, rendering the local motion cues uninformative. Children were required to judge whether there was an intact walker in the mask. The correct response relied on extracting global cues from the intact walker. To prevent children from learning the shape of the walker24, we set target walkers possibly facing one of five equally spaced directions from left to right. Out of the five walkers we used, two faced straight left or right, orthogonal to the viewing directions. Two walked with their bodies oriented at a 45-degree angle to the left or right of the observer. The last one walked towards the observer. Video 4 demonstrated the five facing directions. A total of 30 trials were conducted, consisting of two conditions (intact or scrambled target).

Added in the revised version: Task 3 (BM-General) tested the ability to process general BM cues (local + global cues). In Task 3, the monitor displayed an intact walker (facing either left or right) embedded within a mask (see video 5). The mask used in this task was similar to that in Task 2. Children were required to judge the facing direction of the target walker (left or right). Since the mask and the target walker contained the same local BM cues and the target walker is presented with additional global configuration cues, children could rely on the general BM information (i.e., a combination of local and global cues) BM information to perform the task. Task 3 consisted of 30 trials, with 15 trials for each facing directions. Other parameters of Task 2 and 3 were similar to Task 1. Before each task, children practiced five trials to make sure good understanding. We performed three tasks in a fixed order so that participants were naïve to the nature of local BM cues in Task 1.

In Experiment 2, 42 children with ADHD completed the same procedures as in Experiment 1. Beyond that, their parents should fill in the Social Responsiveness Scale (SRS) to assess social interaction.

Statistics

Earlier version: We carried out ANOVA to examine the difference in BM perception abilities between TD and ADHD children and the Pearson correlation analysis to assess the relationship between the accuracy of each task and the SRS score. General linear models and path analysis were used to explore the roles of potential factors on BM perception. A p value < 0.05 was considered statistically significant.

Revised version: We carried out a two-sample t-test to examine the difference in BM perception abilities between TD and ADHD children47,48 and the Pearson correlation analysis was used to assess the relationship between the accuracy of each task and the SRS score. Additionally, general linear models and path analysis were used to explore the potential factors influencing BM perception. A p-value < 0.05 was considered statistically significant. Path analysis was conducted using AMOS and other analyses were conducted using SPSS.

Results

Children with ADHD exhibit atypical BM perception

Added in the revised version: Figure 3 displays the mean accuracies (ACC) for both the TD and ADHD groups across the three tasks in Experiment 1. We examined the difference in ACC between the TD and ADHD groups for each task using the two-sample t-test. The results of Task 1 showed a significant difference (TD: 0.52 ± 0.13, ADHD: 0.44 ± 0.09, t73 = 3.059, p = 0.003, Cohen’s d = 0.716), indicating that children with ADHD exhibited impaired local BM processing ability. For Task 2 and 3, where children were asked to detect the presence or discriminate the facing direction of the target walker (Task 2 - TD: 0.70 ± 0.12, ADHD: 0.59 ± 0.12, t73 = 3.677, p < 0.001, Cohen’s d = 0.861; Task 3 - TD: 0.79 ± 0.12, ADHD: 0.63 ± 0.17, t73 = 4.702, p < 0.001, Cohen’s d = 1.100).

Removed in the revised version: Figure 3 shows the mean accuracies (ACC) for both groups in the three tasks in Experiment 1. We tested the difference in ACC between the TD and ADHD groups in each task using one-way ANOVA (Gender as a covariate). In Task 1, the group comparison revealed a significant difference (F1,72 = 5.630, p = 0.020, η2 = 0.073). Children are asked to judge the facing direction of the target walker using global BM cues (Task 2) or general BM cues combining global with local BM cues (Task 3). TD group have higher accuracies than the ADHD group (Task 2: F1,72 = 12.600, p < 0.001, η2 = 0.149; Task 3: F1,72 = 19.926, p < 0.001, η2 = 0.217). In all statistical analyses, the main effects of gender were not significant (all p > 0.05).

Removed in the revised version: The mean accuracy of the three tasks.

Earlier version: Typically developing (TD) children had higher accuracies than children with attention deficit hyperactivity disorder (ADHD) in three tasks in Experiment 1. Error bars show standard errors. *p < 0.05; ***p < 0.001

Revised version: The mean accuracy of the three tasks. Typically developing (TD) children had higher accuracies than children with attention deficit hyperactivity disorder (ADHD) in three tasks in Experiment 1. Error bars show standard deviations. ***p < 0.001

Added in the revised version: These findings suggest impaired global and general BM perception in children with ADHD. To further ensure that gender did not influence the results, we conducted a subsampling analysis with balanced data, and the results remained consistent (see Supplementary Information).

Atypical perception of local BM information predicted impaired social interaction in ADHD

Earlier version: Experiment 1 demonstrates that the abilities of BM perception were atypical in ADHD. Our previous study revealed that local BM processing ability is heritable and negatively correlated with autistic traits, especially in the subdimension of communication19. In addition, substantial evidence indicates that children with ADHD often have problems in social interaction. We expected that compromised social interaction in ADHD might be associated with atypical local BM processing. So, we recruited 42 naïve children with ADHD to participate in Experiment 2 to explore this relationship. Their parents or caregivers familiar with children completed SRS (total scores = 54.64 ± 18.42). Thirty-three TD children had SRS scores (total score = 38.64 ± 12.47). The total SRS scores in the TD groups are significantly lower than those in the ADHD group (t = –4.277, p < 0.001). We found that children with higher total score of SRS performed worse in Task 1, i.e., the ability of local BM processing was negatively correlated to SRS total scores (r = –0.264, p = 0.022). Notably, significant correlations were also identified between SRS total score and global BM processing and between SRS total score and general BM processing (global BM perception: r = –0.238, p = 0.039; general BM perception: r = –0.359, p = 0.002).

Revised version: Experiment 1 provided evidence of atypical BM perception in children with ADHD. Previous study revealed that local BM processing ability is heritable and negatively correlated with autistic traits, especially in the subdimension of communication19. In addition, substantial evidence indicates that children with ADHD often have problems in social interaction. We expected that compromised social interaction in ADHD might be associated with atypical local BM processing. To confirm this assumption, we recruited 42 naïve children with ADHD to participate in Experiment 2 and examined the relationship between their social interaction abilities and BM perception. Their parents or caregivers who took care of the children completed SRS. The SRS total score of the ADHD group was higher than that of the TD group (SRS total score - ADHD: 54.64 ± 18.42, TD: 38.64 ± 12.47, t = -4.277, p < 0.001). We found that children with higher total score of SRS performed worse in Task 1, i.e., the ability of local BM processing was negatively correlated to SRS total scores (r = - 0.264, FDR-corrected p = 0.033). We also found significant correlations between SRS total score and global BM processing as well as general BM processing (global BM perception: r = -0.238, FDR-corrected p = 0.039; general BM perception: r = -0.359, FDR-corrected p = 0.006).

Earlier version: In the further subgroup analysis, as seen in Figure 4, the correlation between the SRS total score and the ability to process local cues was only found in the ADHD group (ADHD: r = –0.461, p = 0.002; TD: r = 0.109, p = 0.547), particularly on subscales of social awareness, social cognition, social communication, social motivation (see Table 2 for complete information). However, no correlation was found between SRS scores and global or general BM processing in either ADHD or TD groups (all p > 0.05).

Revised version: In further subgroup analysis, as depicted in Figure 4, the correlation between the SRS total score and the ability to process local cues was only found in the ADHD group (ADHD: r = -0.461, FDR-corrected p = 0.004; TD: r = 0.109, FDR-corrected p = 0.547), particularly on subscales of social awareness, social cognition, social communication, social motivation (see Table 2 for detailed information). However, we did not find significant correlation between SRS total score and global or general BM processing in either the ADHD or TD groups separately (global BM perception - TD: r = -0.020, FDR-corrected p = 0.910, ADHD: r = -0.207, FDR-corrected p = 0.374; general BM perception - TD: r = -0.118, FDR-corrected p = 0.514, ADHD: r = -0.286, FDR-corrected p = 0.134).

Removed in the revised version: Correlations between the accuracy of Task 1 (BM-Local) and SRS total scores.

Earlier version: The correlation between the SRS total score and the ability to process local cues is found in ADHD group. By contrast, this correlation is not found in TD group. The grey shading represents a 95% confidence interval. **p < 0.01.

Revised version: Correlations between response accuracies and SRS total score. The ability to process local cues is significantly correlated with the SRS total score in ADHD group. The grey shading represents a 95% confidence interval. ** FDR- corrected p < 0.01.

Added in the revised version: For an exploration of the distinct relationship between local BM processing and SRS total scores in the ADHD group, we further examined the differences among these correlations (see Supplementary Information). In Task 1, we observed a significantly stronger correlation between SRS total score and the response accuracy of the ADHD group compared to the TD group (p = 0.003). However, this distinction was not identified in Tasks 2 (p = 0.381) or Task 3 (p = 0.455). Additionally, we found that the correlation between SRS total score and the response accuracy of the ADHD group in Task 1 was slightly stronger than in Task 2 (p = 0.066) and Task 3 (p = 0.181). These findings suggested that atypical local BM processing may specifically reflect the impaired social interaction in the ADHD group.

Earlier version: Reasoning intelligence and attention regulated BM processing

Revised version: Global BM processing develops with age and is regulated by reasoning intelligence and attention function

Earlier version: Experiment 2 revealed that local BM processing instead of global BM processing was related to social interaction. Then, which factors would regulate the difference in processing global BM cues? Previous meta-analysis studies of BM perception in ASD found mixed results about the effect of age and IQ on processing BM31, 32. In addition, neural responses related to BM processing are strongly regulated by attention29. Did these factors play critical roles in BM perception in ADHD, a neurodevelopmental disorder characterized by attention deficit? To answer this question, we integrated data from the ADHD group in Experiment 1 and Experiment 2 (n = 80, one child did not complete QB test). Two linear models were built to investigate contributing factors: (a) BM-global = age + gender + FIQ + QbInattention; (b) BM-general = age + gender + FIQ + QbInattention + BM-local + BM-global. We screened factors with the largest contribution to the models by step-wise regression. For model (a), the ability to process global BM cues was enhanced with age (standardized β = 0.251, p = 0.025). For model (b), higher FIQ, particularly on subdimension Perceptual Reasoning (standardized β = 0.271, p = 0.005), and more excellent global BM processing (standardized β = 0.290, p = 0.004) predicted better performance in general BM processing. Furthermore, as children grew up, the ability to probe general BM information increased (standardized β = 0.365, p < 0.001). It is worth noting that QbInattention had a highly negative correlation with Perceptual Reasoning (r = –0.355, p = 0.001) and general BM perception (r = –0.246, p = 0.028). Given the risk of collinearity, we further used path analysis to visualize these relationships post hoc (Figure 5). The result showed sustained attention (i.e., QbInattention) did not predict performance in BM-General (i.e., Task 3) directly. However, a significant indirect path via the ability of Perceptual Reasoning was found. In addition, as children with ADHD grew up, their performance in BM-General (i.e., Task 3) not only improved directly but enhanced via better processing of global BM cues. Moreover, we built model (c) to corroborate that the local BM processing was stable with age and unaffected by attention and IQ: (c) BM-local = age + gender + FIQ + QbInattention. No regressor has remained after step-wise regression.

Revised version: Experiment 2 indicated that local BM processing, as opposed to global BM processing, was more closely associated with social interaction in the ADHD group. To investigate the specificity of this association, we examined the relationship between impairments in different biological processing abilities and factors such as age, IQ, and attention function in the ADHD group. Data from the ADHD group in both Experiment 1 and Experiment 2 were integrated, resulting in a total of 80 ADHD participants (with one child not completing the QB test). Three linear models were built to investigate contributing factors: (a) BM-local = β0 + β1 * age + β2 * gender + β3 * FIQ + β4 * QbInattention; (b) BM-global = β0 + β1 * age + β2 * gender + β3 * FIQ + β4 * QbInattention; (c) BM-general = β0 + β1 * age + β2 * gender + β3 * FIQ + β4 * QbInattention + β5 * BM-local + β6 * BM-global. We screened factors with the largest contribution to the models using step-wise regression. In model (a), no regressor has remained after step-wise regression, suggesting that the local BM processing remained stable with age and was not affected by attention and IQ. In model (b), the ability to process global BM cues was enhanced with age (standardized β1 = 0.251, p = 0.025). In model (c), higher FIQ, particularly on subdimension of Perceptual Reasoning (standardized β3 = 0.271, p = 0.005), and better performance in global BM processing (standardized β6 = 0.290, p = 0.004) predicted better performance in general BM processing. Furthermore, as children grew up, the ability to probe general BM information improved (standardized β1 = 0.365, p < 0.001). It is worth noting that QbInattention showed a strong negative correlation with Perceptual Reasoning (r = -0.355, p = 0.001) and general BM perception (r = -0.246, p = 0.028). Due to the potential issue of collinearity, we employed path analysis post hoc to visualize these relationships (Figure 5). The result indicated sustained attention (i.e., QbInattention) did not directly predict performance in BM-General (i.e., Task 3).

Removed in the revised version: Factors influencing BM perception in ADHD children.

Earlier version: Post hoc path analysis confirmed that the effect of sustained attention to performance in BM-General was totally mediated by Perceptual Reasoning, and the ability of global BM processing partly mediated the effect of age on performance in BM-General.

Revised version: Factors influencing BM perception in ADHD children. Post hoc path analysis confirmed that the effect of sustained attention on performance in BM- General was entirely mediated by Perceptual Reasoning, and the ability of global BM processing partly mediated the effect of age on performance in BM-General.

Added in the revised version: However, a significant indirect path via the ability of Perceptual Reasoning was identified. Furthermore, as children with ADHD grew older, their performance in BM-General (i.e., Task 3) improved, both directly and through enhanced processing of global BM cues.

Earlier version: In addition, we further built three models to explore the effects of Reasoning IQ and age on BM perception in TD children: (a) BM-local = age + gender + FIQ. (b) BM-global = age + gender + FIQ; (c) BM-general = age + gender + FIQ + BM-local +BM-global. No regressor has remained after step-wise regression in model (a). For models (b) and (c), there are positive relations between age and performance (BM-global: standardized β = 0.396, p = 0.017; BM-general: standardized β = 0.330, p = 0.049).

Revised version: We also built three models to further explore the effects of Reasoning IQ and age on BM perception in TD children: (a) BM-local = β0 + β1 * age + β2 * gender + β3 * FIQ. (b) BM-global = β0 + β1 * age + β2 * gender + β3 * FIQ; (c) BM-general = β0 + β1 * age + β2 * gender + β3 * FIQ + β4 * BM-local + β5 * BM-global. In model (a), no regressor remained significant after the step-wise regression. However, in models (b) and (c), we observed positive relationships between age and performance (BM-global: standardized β = 0.396, p = 0.017; BM-general: standardized β = 0.330, p = 0.049).

Added in the revised version: To sum up, our findings suggest that the abilities to perceive global and general BM cues, rather than local BM cues, would improve with age. Age might play different roles in the processing of different BM cues. To further verify this speculation, we examine the differences between the improvements in processing local cues and that in processing global and general cues with age (see Supplementary Information). In the ADHD group, we observed that the abilities to process general BM cues significantly improved with age compared to local cues (p < 0.001). Additionally, there was a trend indicating that the improvement in processing global BM cues with age was greater than that in processing local BM cues (p = 0. 073). However, these patterns were not observed in the TD group (BM-Local vs BM-Global: p = 0. 575; BM-Local vs BM-General: p = 0. 739).

Discussion

Earlier version: Our study contributes several promising findings on atypical biological motion perception in ADHD. Specifically, we found that children with ADHD performed worse in processing local, global and general BM information than TD children. Notably, there was a clear dissociation between local and global BM information processing. The ability to process local BM cues, which is highly heritable, could predict deficits of social interaction in ADHD. By contrast, global BM processing, which is largely shaped by experience, would significantly improve with age. Moreover, the results also unraveled the potential factors related to BM perception: Reasoning Intelligence played a mediating effect on the relationship between attention and general BM perception.

Revised version: Our study contributes several promising findings concerning atypical biological motion perception in ADHD. Specifically, we observe the atypical local and global BM perception in children with ADHD. Notably, a potential dissociation between the processing of local and global BM information is identified. The ability to process local BM cues appears to be linked to the traits of social interaction among children with ADHD. In contrast, global BM processing exhibits an age-related development. Additionally, general BM perception may be affected by factors including attention.

Earlier version: BM perception is a hot topic in the field of visual cognition. One critical reason is that BM naturally contains biological and social properties. BM processing has significant value in successful daily life, particularly in nonverbal communication11, 46 and adaptive behavior29, 47. As an early emerging ability, BM can be discriminated in the newborn baby21. In typical development children, there is a linkage between BM perception and social cognitive abilities48. For example, 12-month-old infants showed social behaviors (i.e., following gaze) elicited by BM displays49. It is, therefore, of the supreme importance of BM to children’s development of social cognition. This “social interpretation” of BM proposed that the ability to process BM may be taken as a measure of difficulties in social interaction19. An example is that poor performance of BM perception is a hallmark of impaired social function in children with ASD10, 50. That implies that atypical BM perception may contribute to the deficits of social interaction in ADHD. Albeit we found a correlation between the SRS total score and the accuracy of each of the tasks in all participants, only a significantly negative correlation between the SRS total score and the accuracy of processing local cues was found in ADHD group in the further subgroup analysis, which provides preliminary evidence for the hypothesis and demonstrates that the local BM processing can predict impaired social interaction in children with ADHD. This finding is consistent with previous studies that revealed a common genetic basis for local BM processing ability and social communication19. In the subgroup analysis, however, we did not find a correlation between social interaction and global BM processing, nor between social interaction and general BM processing. That is because the ability to process global BM cues is greatly affected by environment and general cognition functions19.

Revised version: BM perception is a widely studied topic in the field of visual cognition, due to its inherent biological and social properties. BM processing has significant value in successful daily life, particularly in nonverbal communication11,49 and adaptive behavior29,50. As an early emerging ability, BM can be discriminated in the newborn baby21. In typical development children, there is a clear association between BM perception and social cognitive abilities51. For example, 12-month-old infants exhibit social behaviors (i.e., following gaze) elicited by BM displays52. It is, therefore, of supreme importance of BM to children’s development of social cognition. This ‘social interpretation’ of BM proposed that the difficulties in processing BM may serve as an indicator of impaired social interaction19 This is evident in children with ASD, where impaired BM perception is a characteristic feature linked to impaired social function10,53 .

Added in the revised version: ADHD, a neurodevelopmental disorder, is typically identified in early childhood. Children with ADHD often show impaired social function. We observed atypical BM perception in children with ADHD, potentially contributing to social disfunctions.

Added in the revised version: The proposition that BM perception serves as a prerequisite for attributing intentions to others, facilitating adaptive interaction with other individuals23,54,55, aligns with our discovery of a significant relationship between BM perception performance and the SRS total score. Further subgroup analysis revealed a significant negative correlation between the SRS total score and the accuracy of processing local BM specifically in the ADHD group. Notably, this correlation was stronger in the ADHD group compared to the TD group. This may be due to the narrow range of SRS scores in TD group. Future study should increase sample size to explore the correlation in diverse individuals. To sum up, these findings suggest that local BM processing could be regarded as a distinct hallmark of social cognition in ADHD children10,19.

Added in the revised version: BM perception is considered a multi-level phenomenon56–58. At least in part, processing information of local BM and global BM appears to involve different genetic and neural mechanisms16,19. Using the classical twin method, Wang et al. found that the distinction between local and global BM processing may stem from the dissociated genetic bases. The former, to a great degree, seems to be acquired phylogenetically20,21,59,60, while the latter is primarily obtained through individual development19. The sensitivity to local rather than global BM cues seems to emerge early in life. Visually inexperienced chicks exhibit a spontaneous preference for the BM stimuli of hen, even when the configuration was scrambled20. The same finding was reported in newborns. On the contrary, the ability to process global BM cues rather than local BM cues may be influenced by attention28,29 and shaped by experience24,56. There is evidence of an age-related improvement in global BM, as a previous study indicated that global BM perception was enhanced with age in TD children61,62. This developmental trend also seems to exist in individuals with ASD. The studies in children consistently showed the impairment of global BM processing in children with ASD63,64, whereas no impairment was found in studies with regard to adults with ASD65–67. The performance of BM processing in children with ASD becomes more aligned with TD children as they get older32.

Added in the revised version: Our study contributes new evidence to the understanding of the development of global BM processing. We found that the ability to process global and general BM cues improved significantly with age in both TD and ADHD groups, which imply the processing module for global BM cues tends to be mature with development. In the ADHD group, the improvement in processing general and global BM cues is greater than that in processing local BM cues, while no difference was found in TD group.

Added in the revised version: This may be due to the relatively higher baseline abilities of BM perception in TD children, resulting in a relatively milder improvement. These findings also suggest a dissociation between the development of local and global BM processing. There seems to be an acquisition of ability to process global BM cues, akin to the potential age-related improvements observed in certain aspects of social cognition deficits among individuals with ADHD5, whereas local BM may be considered an intrinsic trait19. It is worth noting that the ability to process global BM cues is positively correlated with the performance of processing general BM cues in ADHD group, while no such correlation was not found in TD group. This suggests that TD children are able to extract and integrate both local and global cues, whereas children with

Added in the revised version: ADHD may rely more on global BM cues to judge the facing direction of the walker when presented with both local and global BM cues, which correspond to a hierarchical model56. Once a living creature is detected, an agent (i.e., is it a human?) can be recognized by a coherent, articulated body structure perceptually organized from its motions (i.e., local BM cues)68. This process involves a top-down processing and probably requires attention28,69, especially in the presence of competing information29. Our findings are consistent with previous research on the cortical processing of BM70, as we found the severity of inattention in the children with ADHD was negatively correlated to their performance of BM processing. More future works are required to verify these conjectures.

Removed in the revised version: BM perception is considered a multi-level phenomenon51–53. At least in part, processing information of local BM and global BM are different16, 19. Our results also demonstrate such a dissociation. Local motion processing is stable with age and not susceptible to attention and IQ, which shows a phylogenetically highly preserved mechanism20, 21, 54, 55 and benefits of life detection22. On the contrary, a development tendency seems to exist in global BM processing that tends to be improved with age. A previous study indicated BM perception was enhanced with age in TD children56. This variation is also observed in children with ASD, and some developmental overlaps were found in the ability to process visual motion information57. Performance of BM processing within ASD becomes more aligned with TD children, and the discrepancy between the two groups alleviates with the increase in age32. With regard to ADHD, deficits of social cognition are likely to improve with age in most individuals 5. Our study further verified this point in BM perception. It is worth noting that the ability of global BM processing positively correlated with the performance of processing general BM cues, i.e., children with ADHD are apt to catch global BM cues to judge the facing direction of the walker when displayed with local and global BM cues, which correspond to a hierarchical model51. Once a living creature is detected, an agent (i.e., is it a human?) can be recognized by a coherent, articulated body structure perceptually organized from its motions (i.e., local BM cues)58, which is a top-down processing and probably requires attention28, 59, especially in the presence of competing information29. In addition, cortical processing of BM is strongly modulated by attention60. Therefore, as a disorder characterized by deficits of sustained attention, we found its severity of inattention was negatively related to the performance of BM processing.

Earlier version: Interestingly, a compensation strategy was found in children with impaired BM perception. No difficulty of BM recognition was found in ASD individuals with high IQ61, but children with ASD showed weaker adaptation effects for biological motion than TD children62. One possibility is that individuals with high IQ and impaired BM perception can better create or reason reliable strategies used for BM recognition63, 64. The current study supported this assumption. Children with higher IQ, particularly Perceptual Reasoning, have better performance. Because of the effect of attention deficits on Perceptual Reasoning, the performance within ADHD will not become aligned with that in TD children.

Revised version: Interestingly, children with impaired BM perception may employ a compensation strategy71. Previous study found no impairment in BM recognition in autistic individuals with high IQ71, but children with ASD exhibited weaker adaptation effects for biological motion compared to TD children72. One possibility is that individuals with high IQ and impaired BM perception are able to develop or employ reliable strategies used for BM recognition, compensating for the lack of intuitive social perceptual processing73,74. The current study supports this assumption, as children with higher IQ, particularly in Perceptual Reasoning, demonstrated better performance. Due to the impact of attention deficits on Perceptual Reasoning, the performance of children with ADHD did not align with that of TD children.

Earlier version: Overall, our study first revealed two atypical fundamental abilities underlying BM perception in children with ADHD, which showed a clear dissociation. Significantly, anomalous local BM processing predicts impaired social interaction in ADHD. Moreover, these results unraveled the potential contributions of age, IQ and attention to processing BM information. These findings also shed new light on the direction of future studies, which should explore the performance of more advanced BM processing in children with ADHD, such as emotion and identity recognition in BM tasks. In addition, it is necessary to delineate the profiles for neural processing of biological motion in ADHD. A comparative study is warranted to investigate differences between ADHD and ASD, which is essential for finding the common neuropsychological traits and biomarkers of social cognition impairments.

Revised version: Overall, our study first revealed two distinct and atypical fundamental abilities underlying BM perception in children with ADHD, indicating a potential dissociation of local and global BM processing. Notably, anomalous local BM processing could predict impaired social interaction in ADHD. Moreover, these results unraveled the potential contributions of age, IQ and attention to the process of BM information.

Added in the revised version: These findings also shed new light on the direction of future studies. Firstly, more studies are required to confirm the dissociation between the two fundamental BM processing abilities, which will contribute to understanding the respective neural mechanisms underlying the two BM processing. Secondly, exploring the performance of more advanced BM processing in children with ADHD, such as emotion and identity recognition in BM tasks, is necessary to delineate the neural profiles for the processing of biological motion in ADHD. Finally, a comparative study between ADHD and ASD is warranted to identify common neuropsychological traits and biomarkers of social cognition impairments.

Added in the revised version: Supporting information